Nebulised dornase alfa reduces inflammation and improves clinical outcomes in severe COVID-19: a randomised clinical trial

Author Response

The following is the authors’ response to the original reviews.

Reviewer 1:

Public review:

In this study, Porter et al report on outcomes from a small, open-label, pilot randomized clinical trial comparing dornase-alfa to the best available care in patients hospitalized with COVID-19 pneumonia. As the number of randomized participants is small, investigators describe also a contemporary cohort of controls and the study concludes about a decrease of inflammation (reflected by CRP levels) aJer 7 days of treatment but no other statistically significant clinical benefit.

Suggestions to the authors:

• The RCT does not follow CONSORT statement and reporting guidelines

We thank you for this suggestion and have now amended the order and content of the manuscript to follow the CONSORT statement as closely as possible.

• The authors have chosen a primary outcome that cannot be at least considered as clinically relevant or interesting. AJer 3 years of the pandemic with so much research, why investigate if a drug reduces CRP levels as we already have marketed drugs that provide beneficial clinical outcomes such as dexamethasone, anakinra, tocilizumab and baricitinib.

We thank the reviewer for bringing up this central topic. The answer to this question has both a historical and practical component. This trial was initiated in June of 2020 and was completed in June of 2021. At that time there were no known treatments for the severe immune pathology of COVID19 pneumonia. In June 2020, dexamethasone data came out and we incorporated dexamethasone into the study design. It took much longer for all other anti-inflammatories to be tested. Hence, our decision to trial an approved endonuclease was based purely on basic science work on the pathogenic role of cell-free chromatin and NETs in murine sepsis and flu models and the ability of DNase I to clear them and reduce pathology in these animal models. In addition, evidence for the presence of cell-free chromatin components in COVID-19 patient plasma had already been communicated in a pre-print. Finally, several studies had reported the anti-inflammatory effects of dornase treatment in CF patients. Hence there was a strong case for a cheap, safe, pulmonary noninvasive treatment that could be self-administered outside the clinical se]ng.

The Identification of novel/repurposed treatments effective for COVID-19 were hampered by patient recruitment to competing studies during a pandemic. This resulted in small studies with inconclusive or contrary findings. In general, effective treatments were only picked up in very large RCTs. For example, demonstrating dexamethasone as effective in COVID-19 required recruitment of 6,425 patients into the RECOVERY study. Multiple trials with anti-IL-6 gave conflicting evidence until RECOVERY recruited 4116 adults with COVID-19 (n=2022, tocilizumab and 2094, control) similar for Baracitinib (4,148 randomised to treatment and 4,008 to standard care). Anakinra is approved for patients with elevated suPAR, based on data from one randomized clinical trial of 594 patients, of whom 405 had active treatment (PMID: 34625750). However, a systematic review analysing over 1,627 patients (anakinra 888, control 739) with COVID-19 showed no benefit (PMID: 36841793). Regarding the choice of the primary endpoint, there is a wealth of clinical evidence to support the relevance of CRP as a prognostic marker for COVID-19 pneumonia patients and it is a standard diagnostic and prognostic clinical parameter in infectious disease wards. This choice in March 2020 was supported by evidence of the prognostic value of IL-6; CRP is a surrogate of IL-6. We also provide our own data from a large study of severe COVID-19 pneumonia in figure 1, showing how well CRP correlates with survival.

In summary, our data suggest that Dornase yields an anti-inflammatory effect that is comparable or potentially superior to cytokine-blocking monotherapies at a fraction of the cost and potentially without the additional adverse effects such as the increase for co-infections.

We now provide additional justification on these points in the introduction on pg.4 as follows:

“The trial was ini.ated in June 2020 and was completed in September of 2021. At the start of the trial only dexamethasone had been proven to benefit hospitalized COVID-19 pneumonia pa.ents and was thus included in both arms of the trial. To increase the chance of reaching significance under challenging constraints in pa.ent access, we opted to increase our sample size by using a combina.on of randomized individuals and available CRP data from matched contemporary controls (CC) hospitalized at UCL but not recruited to a trial. These approaches demonstrated that when combined with dexamethasone, nebulized DNase treatment was an effec.ve an.-inflammatory treatment in randomized individuals with or without the implementa.on of CC data.”

We also added the following explanation in the discussion on pg. 16:

“Our study design offered a solution to the early screening of compounds for inclusion in larger platform trials. The study took advantage of frequent repeated measures of quantifiable CRP in each patient, to allow a smaller sample size to determine efficacy/futility than if powered on clinical outcomes. We applied a CRP-based approach that was similar to the CATALYST and ATTRACT studies. CATALYST showed in much smaller groups (usual care, 54, namilumab, 57 and infliximab, 35) that namilumab that is an antibody that blocks the cytokine GM-CSF reduced CRP even in participants treated with dexamethasone whereas infliximab that targets TNF-α had no significant effect on CRP. This led to a suggestion that namilumab should be considered as an agent to be prioritised for further investigation in the RECOVERY trial. A direct comparison of our results with CATALYST is difficult due to the different nature of the modelling employed in the two studies. However, in general Dornase alfa exhibited comparable significance in the reduction in CRP compared to standard of care as described for namilumab at a fraction of the cost. Furthermore, endonuclease therapies may prove superior to cytokine blocking monotherapies, as they are unlikely to increase the risk for microbial co-infections that have been reported for antibody therapies that neutralize cytokines that are critical for immune defence such as IL-1β, IL-6 or GM-CSF. “

• Please provide in Methods the timeframe for the investigation of the primary endpoint

This information is provided in the analysis on pg. 8:

“The primary outcome was the least square (LS) mean CRP up to 7 days or at hospital discharge whichever was sooner.”

• Why day 35 was chosen for the read-out of the endpointt?

We now state on pg. 8 that “Day 35 was chosen as being likely to include most early mortality due to COVID-19 being 4 weeks after completion of a week of treatment. ( i.e. d7 of treatment +28 (4 x 7 days))”

• The authors performed an RCT but in parallel chose to compare also controls. They should explain their rationale as this is not usual. I am not very enthusiastic to see mixed results like Figures 2c and 2d.

We initially aimed at a fully randomized trial. However, the swiJ implementation of trial prioritization strategies towards large and pre-established trial plamorms in the UK made the recruitment COVID19 patients to small studies extremely challenging. Thus, we struggled to gain access to patients. Our power calculations suggested that a mixed trial with randomized and contemporary controls was the best way forward under these restrictions in patient access that could provide sufficient power.

That being said, we also provide the primary endpoint (CRP) results in Fig. 3B as well as the results for the length of hospitalization (Fig. S3D) for the randomized subjects only.

• Analysis is performed in mITT; this is a major limitation. The authors should provide at least ITT results. And they should describe in the main manuscript why they chose mITT analysis.

We apologize if this point was confusing. We performed the analysis on the ITT as defined in our SAP: “The primary analysis population will be all evaluable patients randomised to BAC + dornase alfa or BAC only who have at least one post-baseline CRP measurement, as well as matched historical comparators.”

We understand that the reason this might be mistaken as an mITT is because the N in the ITT (39) doesn’t match the number randomised and because we had stated on pg. 8 that “ Efficacy assessments of primary and secondary outcomes in the modified inten.on-to-treat popula.on were performed.”

However, we did randomise 41 participants, but:

One participant in the DA arm never received treatment. The individual withdrew consent and was replaced. We also have no CRP data for this participant in the database, so they were unevaluable, and we couldn’t include them in the baseline table even if we wanted to. In addition, 1 participant in BAC only had a baseline CRP measurement available. Hence not evaluable as we only have a baseline CRP measurement for this participant.

We have corrected the confusing statement on pg. 8 and added an additional explanation.

“Efficacy assessments of primary and secondary outcomes in the inten.on-to-treat (ITT) popula.on were performed on all randomised par.cipants who had received at least one dose of dornase alfa if randomized to treatment. For full details see Sta.s.cal Analysis Plan. The ITT was adjusted to mi.gate the following protocol viola.ons where one par.cipant in the BAC arm and one in the DA arm withdrew before they received treatment and provided only a baseline CRP measurement available. The par.cipant in the DA arm was replaced with an addi.onal recruited pa.ent. Exploratory endpoints were only available in randomised par.cipants and not in the CC. In this case, a post hoc within group analysis was conducted to compare baseline and post-baseline measurements.”

• It is also not usual to exclude patients from analysis because investigators just do not have serial measurements. This is lost to follow up and investigators should have pre-decided what to do with lost-to-follow-up.

Our protocol pre-specified that the primary analysis population should have at least one postbaseline CRP measurement (pg. 13 of protocol). The patient that was excluded was one that initially joined the trial but withdrew consent after the first treatment but before the first post-treatment blood sample could be drawn. Hence, the pre-treatment CRP of this patient alone provided no useful information.

• In Table 1 I would like to see all randomized patients (n=39), which is missing. There are also baseline characteristics that are missing, like which other treatments as BAT received by those patients except for dexamethasone.

Table 1 includes all 39 patients plus 60 CCs.

Table 2 shows additional treatments given for COVID-19 as part of BAC.

• In the first paragraph of clinical outcomes, the authors refer to a cohort that is not previously introduced in the manuscript. This is confusing. And I do not understand why this analysis is performed in the context of this RCT although I understand its pilot nature.

One of the main criticisms we have encountered in this study has been the choice of the primary endpoint. The best way respond to these questions was to provide data to support the prognostic relevance of CRP in COVID-19 pneumonia from a separate independent study where no other treatments such as dexamethasone, anakinra or anti-IL6 therapies were administered. We think this is very useful analysis and provides essential context for the trial and the choice of the primary endpoint, indicating that CRP has good enough resolution to predict clinical outcomes.

• Propensity-score selected contemporary controls may introduce bias in favor of the primary study analysis, since controls are already adjusted for age, sex and comorbidities.

The contemporary controls were selected to best match the characteristics of the randomized patients including that the first CRP measurement upon admission surpassed the trial threshold, so we do not see how this selection process introduces biases, as it was blinded with regards to the course of the CRP measurements. Given that this was a small trial, matching for baseline characteristics is necessary to minimize confounding effects.

• The authors do not clearly present numerically survivors and non-survivors at day 34, even though this is one of the main secondary outcomes.

We now provide the mortality numbers in the following paragraph on pg. 13.

“Over 35 days follow up, 1 person in the BAC + dornase-alfa group died, compared to 8 in the BAC group. The hazard ra.o observed in the Cox propor.onal hazards model (95% CI) was 0.47 (0.06, 3.86), which es.mates that throughout 35 days follow-up, there was a 53% reduced chance of death at any given .mepoint in the BAC + dornase-alfa group compared to the BAC group, though the confidence intervals are wide due to a small number of events. The p-value from a log-rank test was 0.460, which does not reach sta.s.cal significance at an alpha of 0.05.”

• It is unclear why another cohort (Berlin) was used to associate CRP with mortality. CRP association with mortality should (also) be performed within the current study.

As we explained above, the Berlin cohort CRP data serve to substantiate the relevance of CRP as a primary endpoint in a cohort that experienced sufficient mortality as this cohort did not receive any approved anti-inflammatory therapy. Mortality in our COVASE trial was minimal, since all patients were on dexamethasone and did not reach the highest severity grade, since we opted to treat patients before they deteriorated further. The overall mortality was 8% across all arms of our study, which does not provide enough events for mortality measurements. In contrast the Berlin cohort did not receive dexamethasone and all patients had reached a WHO severity grade 7 category with mortality at 30%.

My other concerns are:

• This report is about an RCT and the authors should follow the CONSORT reporting guidelines. Please amend the manuscript and Figure 1b accordingly and provide a CONSORT checklist.

We now provide a CONSORT checklist and have amended the CONSORT diagram accordingly.

• Please provide in brief the exclusion criteria in the main manuscript

We have now included the exclusion criteria in the manuscript on pg. 6.

“1.1.1 Exclusion criteria

1. Females who are pregnant, planning pregnancy or breasmeeding

2. Concurrent and/or recent involvement in other research or use of another experimental inves.ga.onal medicinal product that is likely to interfere with the study medica.on within (specify .me period e.g. last 3 months) of study enrolment 3. Serious condi.on mee.ng one of the following:

a. Respiratory distress with respiratory rate >=40 breaths/min

b. oxygen satura.on<=93% on high-flow oxygen

1. Require mechanical invasive or non-invasive ven.la.on at screening

2. Concurrent severe respiratory disease such as asthma, COPD and/or ILD

3. Any major disorder that in the opinion of the Inves.gator would interfere with the evalua.on of the results or cons.tute a health risk for the trial par.cipant

4. Terminal disease and life expectancy <12 months without COVID-19

5. Known allergies to dornase alfa and excipients

6. Par.cipants who are unable to inhale or exhale orally throughout the en.re nebulisa.on period So briefly Pa.ents were excluded if they were:

7. pregnant, planning pregnancy or breasmeeding 2. Serious condi.on mee.ng one of the following:

a. Respiratory distress with respiratory rate >=40 breaths/min

b. oxygen satura.on<=93% on high-flow oxygen

1. Require ven.la.on at screening

2. Concurrent severe respiratory disease such as asthma, COPD and/or ILD

3. Terminal disease and life expectancy <12 months without COVID-19

4. Known allergies to dornase alfa and excipients

5. Par.cipants who are unable to inhale or exhale orally throughout the en.re nebulisa.on period”

• "The final trial visit occurred at day 35." "Analysis included mortality at day 35". I am not sure I understand why. In clinicaltrials.gov all endpoints are meant to be studies at day 7 except for mortality rate day 28. Why day 35 was chosen? Please be consistent.

Thank you for identifying this inconsistency. We have amended the record on clinicaltrials.gov to read ‘’the time to event data was censored at 28 days post last dose (up to d35) for the randomised participants and at the date of the last electronic record for the CC.”

• Please provide in Methods the timeframe for the investigation of the primary endpoint

• The authors performed an RCT but in parallel chose to compare also controls. They should explain their rationale as this is not usual. I am not very enthusiastic to see mixed results like Figures 2c and 2d.

• Analysis is performed in mITT; this is a major limitation. The authors should provide at least ITT results. And they should describe in the main manuscript why they chose mITT analysis.

• It is also not usual to exclude patients from analysis because investigators just do not have serial measurements. This is lost to follow up and investigators should have pre-decided what to do with lost-to-follow-up.

• Figure 1b as in CONSORT statement, please provide reasons why screened patients were not enrolled.

• In Table 1 I would like to see all randomized patients (n=39), which is missing. There are also baseline characteristics that are missing, like which other treatment as BAT received those patients except for dexamethasone.

• In the first paragraph of clinical outcomes, the authors refer to a cohort that is not previously introduced in the manuscript. This is confusing. And I do not understand why this analysis is performed in the context of this RCT although I understand its pilot nature.

• In Figure 2 the authors draw results about ITT although in methods describe that they performed an mITT analysis. Please be consistent.

Please see answers provided to these queries above.

Reviewer #2 (Recommendations For The Authors):

1. Suppl Figure 2B would be more informative if presented as a Table with N of patients with per day sampling

We now provide the primary end point daily sampling table in Table 3.

1. The numbers at risk should figure under the KM curves

The numbers at risk for figures 1E, 2C, 2D have been added as graphs either in the main figures or in the supplement.

1. HD in Supplementary figure 3 should be explained

We apologize for this omission. We now provide a description for the healthy donor samples that we used in the cell-free DNA measurements in figure S3B on pg. 14:

“Compared to the plasma of anonymized healthy donors volunteers at the Francis Crick ins.tute (HD), plasma cf-DNA levels were elevated in both BAC and DA-treated COVASE par.cipants.

1. Presentation is inappropriate for Table S4

We thank the reviewer for pointing this issue. We have now formaxed Table S4 to be consistent with all other tables.

Changes from the earlier version to the revised version

Abstract

Abstract

Background

Added in the revised version: Cell-free (cf)-DNA, from cellular sources, including Neutrophil Extracellular Traps (NETs), is found in the circulation of COVID-19 patients and may contribute to immune dysregulation. This study determined whether pulmonary administration of the endonuclease, dornase alfa, reduced systemic inflammation by degrading local and systemic cf-DNA.

Removed in the revised version: SARS-CoV2 infection causes severe, life-threatening pneumonia. Hyper-inflammation, coagulopathy and lymphopenia are associated with pathology and poor outcomes in these patients. Cell-free (cf) DNA is prominent in COVID-19 patients, amplifies inflammation and promotes coagulopathy and immune dysfunction. We hypothesized that cf-DNA clearance by nebulised dornase alfa may reduce inflammation and improve disease outcomes. Here, we evaluated the efficacy of nebulized dornase alfa in patients hospitalised with severe COVID-19 pneumonia.

Methods

Added in the revised version: Eligible patients were randomized (3:1) to receive twice-daily nebulised dornase alfa in addition to best available care (BAC) or BAC alone for seven days or until discharge. A 2:1 ratio of matched contemporary controls (CC) provided additional comparators. The primary endpoint was improvement in C-reactive protein (CRP) over time, analysed using a repeated-measures mixed model, adjusted for baseline factors.

Removed in the revised version: In this randomised controlled single-centre phase 2 proof-of-concept trial, we recruited adult patients admitted to hospital that exhibited stable oxygen saturation (≥94%) on supplementary oxygen and a C-reactive protein (CRP) level ≥30mg/L post dexamethasone treatment. Participants were randomized at a 3:1 ratio to receive twice-daily nebulised dornase alfa in addition to best available care (BAC) or BAC alone for seven days or until hospital discharge. A 2:1 ratio of historical controls to treated individuals (HC, 2:1) were included as the primary endpoint comparators. The primary outcome was a reduction in systemic inflammation measured by blood CRP levels over 7 days post-randomisation, or to discharge if sooner. Secondary and exploratory outcomes included time to discharge, time on oxygen, D-dimer levels, lymphocyte counts and levels of circulating cf-DNA.

Results

Earlier version: We screened 75 patients and enrolled 39 participants out of which 30 in dornase alfa arm, and 9 in BAC group. We also matched the recruited patients in the treated group (N=30) to historical controls in the BAC group (N=60). For the the primary outcome, 30 patients in the dornase alfa were compared to 69 patients in the BAC group. Dornase alfa treatment reduced CRP by 33% compared to the BAC group at 7-days (P=0.01). The dornase alfa group least squares mean CRP was 23.23 mg/L (95% CI 17.71 to 30.46) and the BAC group 34.82 mg/L (95% CI 28.55 to 42.47). A significant difference was also observed when only randomised participants were compared. Furthermore, compared to the BAC group, the chance of live discharge was increased by 63% in the dornase alfa group (HR 1.63, 95% CI 1.01 to 2.61, P=0.03), lymphocyte counts were improved (least-square mean: 1.08 vs 0.87, P=0.02) and markers of coagulopathy such as D-dimer were diminished (least-square mean: 570.78 vs 1656.96μg/mL, P=0.004). Moreover, the dornase alfa group exhibited lower circulating cf-DNA levels that correlated with CRP changes over the course of treatment. No differences were recorded in the rates and length of stay in the ICU or the time on oxygen between the groups. Dornase alfa was well-tolerated with no serious adverse events reported.

Revised version: Between June 2020-October 2021 we recruited 39 evaluable patients: 30 randomised to dornase alfa; 9 randomised to BAC; with 60 CC. Dornase alfa reduced CRP by 33% compared to BAC. Least squares (LS) mean post-dexamethasone CRP fell from 101.9mg/L to 23.23 mg/L in the BAC+ dornase alfa group versus a fall from 99.5mg/L to 34.82 mg/L in the BAC group at 7 days; P=0.01. This effect of dornase alfa on CRP was confirmed with subgroup and sensitivity analyses. Dornase alfa increased the chance of live discharge by 63% (HR 1.63, 95% CI 1.01 to 2.61, P=0.03), increased lymphocyte counts (LSmean 1.08 vs 0.87, P=0.02) and reduced circulating cf-DNA and the coagulopathy marker D-dimer (LS mean: 570.78 vs 1656.96 μg/mL, P=0.004). Dornase alfa was well-tolerated.

Conclusions

Added in the revised version: We provide proof-of-concept evidence that dornase alfa reduces pathogenic inflammation in hospitalised patients with COVID-19 pneumonia.

Removed in the revised version: In this proof-of-concept study in patients with severe COVID-19 pneumonia, treatment with nebulised dornase alfa resulted in a significant reduction in inflammation, markers of immune pathology and time to discharge. The effectiveness of dornase alfa in patients with acute respiratory infection and inflammation should be investigated further in larger trials.

Introduction

Added in the revised version: SARS-CoV-2 pneumonia can lead to hyperinflammation, coagulopathy, respiratory failure, and death (Siddiqi and Mehra, 2020; Zhou et al., 2020), even with less pathogenic variants and in immunised patients (2000). In severe COVID-19 neutrophil activation drives neutrophil extracellular trap (NET) formation (Radermecker et al., 2020; Zuo et al., 2020). NETs, are composed of DNA-histone complexes and have been associated with coagulopathy and endothelial dysfunction (Papayannopoulos, 2018). Moreover, extracellular histones promote inflammation, immune dysfunction and lethality in sepsis (Ioannou et al., 2022; Tsourouktsoglou et al., 2020; Xu et al., 2011; Xu et al., 2009). Digestion of chromatin DNA by endonucleases suppresses the proinflammatory activity of histones and enables their clearance from the circulation (Ioannou et al., 2022; Tsourouktsoglou et al., 2020). Consistently, DNAse I treatment reduces pathology in murine pulmonary viral infections (Cortjens et al., 2018; Pillai et al., 2016). Endogenous DNAse activity and NET clearance capacity are defective in severe COVID-19 pneumonia (Aramburu et al., 2022b), and the extent of these defects correlates with mortality. Hence, supplementation with exogenous DNases may facilitate extracellular chromatin degradation and increase survival.

Added in the revised version: Pulmozyme, dornase alfa, is a recombinant human DNAase approved since 1993 for patients with cystic fibrosis (CF)(Konstan and Ratjen, 2012; Lazarus and Wagener†, 2019). Dornase alfa solubilizes NETs, reduces inflammation and improves pulmonary function in chronic and acute exacerbations of CF(Konstan and Ratjen, 2012; Papayannopoulos et al., 2011). Pulmozyme is safe and well-tolerated in children and adults with CF at doses up to 10mg BD. Nebulized pulmozyme does not increase circulating endonucleases. Hence, it remained unclear whether the degradation of extracellular chromatin in the lungs would reduce circulating chromatin and systemic hyper-inflammation.

Added in the revised version: To probe the effectiveness of DNase I as an anti-inflammatory therapy in severe COVID-19 pneumonia, we measured as our primary endpoint C-reactive protein (CRP). CRP is a reliable marker of systemic inflammation in severe infection(McArdle et al., 2004). In patients infected with COVID-19, a high CRP of ≥40 mg/L carried a poor prognosis with a lower cut-off of ≥35 mg/L in the elderly (Villoteau et al., 2021), particularly when combined with lymphopenia and coagulopathy marked by elevated D-dimer (Fisher et al., 2022; Liu et al., 2020; Luo et al., 2020; Smilowitz et al., 2021; Tornling et al., 2021; Ullah et al., 2020). Poor outcomes are associated with an early rise (Mueller et al., 2020), higher peak and delayed reduction in CRP (Cui et al., 2021). Moreover, elevated IL-6 and CRP can predict the need for mechanical ventilation(Herold et al., 2020). In addition, CRP correlated well with levels of DNAse and NET degradation activity in COVID-19 patient plasma suggesting a link between CRP and the capacity to degrade cf-DNA in the circulation (Aramburu et al., 2022b). CRP became a reliable primary endpoint in subsequent trials demonstrating the effectiveness of other systemic therapies such as namilumab or infliximab that target inflammatory cytokines (Fisher et al., 2022).

Added in the revised version: In this proof-of-concept study we evaluated the effect of dornase alfa on inflammation, as measured by the impact on circulating CRP in patients with COVID-19, compared with best available care (BAC). The trial was initiated in June 2020 and was completed in October 2021. At the start of the trial only dexamethasone had been proven to benefit hospitalized COVID-19 pneumonia patients and was thus included in both arms of the trial. To increase the chance of reaching significance under challenging constraints in patient access, we opted to increase our sample size by using a combination of randomized individuals and available CRP data from matched contemporary controls (CC) hospitalized at UCL but not recruited to a trial. These approaches demonstrated that when combined with dexamethasone, nebulized DNase treatment was an effective anti-inflammatory treatment in randomized individuals with or without the implementation of CC data.

Added in the revised version: Methods

Added in the revised version: Sponsor and location

Added in the revised version: The trial was sponsored by University College London (UCL) and carried out at University College London Hospital UCLH with ethical (REC: 20/SC/0197, Protocol: 132333, RAS ID:283091) and UK MHRA approvals. All randomized participants provided informed consent. Consent for CC was covered by Health Service (Control of Participant Information) Regulations 2002. Safety and data integrity were overseen by the Trial Monitoring Group and Data Monitoring Committee. All data was collected at UCLH. Additional CRP data is reported from the Pa-COVID-19 study, Charité Universitätsmedizin Berlin with ethical approvaL, Berlin (EA2/066/20). Both Pa-COVID-19 and COVASE studies were carried out according to the Declaration of Helsinki and the principles of Good Clinical Practice (ICH 1996).

Removed in the revised version: Severe SARS-CoV-2 infection is associated with severe pneumonia, hyperinflammation, coagulopathy, respiratory failure and death1,2. Hyperinflammation and a dysregulated immune response play critical roles in the pathophysiology of COVID-19 pneumonia.

Removed in the revised version: In severe COVID-19 pneumonia excessive neutrophil activation drives neutrophil extracellular trap (NET) formation3,4. NETs are composed of DNA, histones and other components and contribute to organ damage by promoting coagulopathy and endothelial dysfunction5. Moreover, extracellular histones promote inflammation and lethality during sepsis6-9. Consistently, DNAse treatment reduces pathology in murine pulmonary viral infections10,11. A treatment that enhances cf-DNA clearance is likely to reduce hyper-inflammation and coagulopathy in severe COVID-19 pneumonia and save lives. Pulmozyme®, dornase alfa, is a recombinant human DNAase approved since 1993 as a nebulised treatment for patients with cystic fibrosis (CF)12. Dornase alfa solubilizes NETs, reduces inflammation and improves pulmonary function in chronic and acute exacerbations of CF13,14. Pulmozyme is safe and well-tolerated in children and adults with CF at doses up to 10mg BD.

Removed in the revised version: C-reactive protein (CRP) is a prognostic marker that correlates with clinical symptoms, hyper-inflammation and response to therapy. CRP correlates is a widely used marker of systemic inflammation. In blood the normal concentration of CRP is <0.8 mg/L but it rises rapidly peaking at 48h from disease onset. A raised CRP above 40 mg/L distinguishes patients with severe COVID-19 from those with mild illness15. We chose CRP as our primary endpoint as it is easily measured and is highly responsive to treatments that reduce inflammation.

Removed in the revised version: Here we report the results of the COVASE study, a Phase IIa trial that evaluated the safety and the efficacy of nebulised dornase alfa in reducing hyperinflammation in hospitalised patients with severe COVID-19 pneumonia, as measured by C-reactive protein (CRP) and clinical endpoints.

Removed in the revised version: Materials and Methods

Removed in the revised version: Trial oversight

Removed in the revised version: The trial was sponsored by University College London and carried out at University College London Hospitals (UCLH). The trial protocol and the statistical analysis plan are available in the Appendix. The COVASE trial was reviewed by the South Central - Hampshire B Research Ethics Committee Level 3 Block B Whitefriars Lewins Mead Bristol BS1 2NT, chaired by Professor Vincenzo Libri (REC reference: 20/SC/0197, Protocol number: 132333, RAS project ID:283091) and the United Kingdom Medicines and Healthcare products Regulatory Agency (MHRA). The trial was conducted in accordance with the principles of the Declaration of Helsinki and the ethical guidelines of the Council for International Organizations of Medical Sciences, applicable International Council for Harmonisation Good Clinical Practice guidelines, and applicable laws and regulations. All the randomized participants provided written informed consent. Consent for the historical controls was covered by the Health Service (Control of Participant Information) Regulations 2002 that allows the processing of Confidential Participant Information (CPI) for specific purposes. Regulation 3 provides for the processing of CPI in relation to communicable diseases and other threats to public health and in particular, allows the Secretary of State to require organisations to process CPI for purposes related to communicable diseases. The COVID-19 pandemic is covered by this legislation which allows a range of purposes related to diagnosing, managing, and controlling the spread of COVID-19. The sponsors designed the trial in collaboration with the investigators at the Francis Crick institute, Exploristics and Target to Treatment Consulting. The sponsors and trial investigators participated in data collection, analysis, and interpretation. The authors made the decision to submit the manuscript for publication and vouch for the accuracy and completeness of the data presented and for the fidelity of the trial to the protocol. The COVASE study was registered on clinicaltrials.gov identifier: NCT04359654. Safety and data integrity were regularly reviewed by the Trial Monitoring Group and Data Monitoring Committee.

Removed in the revised version: Participant population

Removed in the revised version: Adults (≥18 years of age) admitted to UCLH with confirmed SARS-Cov2 infection by reverse-transcriptase–polymerase-chain-reaction and radiologically confirmed COVID-19 pneumonia on chest radiograph or CT-scan; an oxygen saturation below 94% requiring the use of supplemental oxygen; and evidence of hyperinflammation (C-reactive protein [CRP] ≥ 30 mg/L, after administration of dexamethasone) were eligible16. Participants that had unstable ventilatory requirements or needed intubation and ventilation within 24 hours of admission were excluded. Full inclusion and exclusion criteria can be found in the Protocol: Supplementary Appendix 1.

Trial design

Earlier version: The COVASE trial was a single-site, randomised, controlled, parallel, open-label investigation of the effect of dornase alfa on hyperinflammation in hospitalised participants with COVID-19. Screening was performed within 24 hours prior to the administration of dornase alfa (Figure 1A). Eligible, consented participants were randomly assigned at a 3:1 ratio with the use of a closed envelope system to receive either best available care (BAC) plus nebulised dornase alfa or BAC alone. On Day 1 the first baseline sample was collected prior to dornase alfa administration. Thereafter, from Day 1 to Day 7 of the trial participants randomised to the active arm received 2.5mg BD of nebulised dornase alfa in addition to BAC. In all cases BAC included dexamethasone (6 mg/day) for 10 days or until hospital discharge, whichever was shorter as per the RECOVERY trial17. Participants received additional treatments at the discretion of their physicians. The primary analysis was performed on samples up to Day 7. The final trial visit occurred at day 35. In addition to participants randomised to receive BAC alone, for every COVASE participant randomised to active treatment, 2 matched historical controls were included. Historical controls (HCs) had been admitted to UCLH and treated with the same BAC treatment strategy, including dexamethasone. HCs were identified from a database of >600 patients, and were matched for age, gender, BMI, comorbidities and CRP (either pre-or post-dexamethasone to provide 2 HC populations). In most cases the patients in the historical controls were admitted concomitantly with the participants recruited to the trial. Additional details regarding the trial design are provided in the protocol (SA1).

Revised version: The COVASE trial was a single-site, randomised, controlled, parallel, open-label investigation. The primary endpoint was change in CRP. Screening was performed within 24 hours prior to the administration of dornase alfa (Figure 1A). Eligible, consented participants were randomly assigned (3:1) to either BAC plus nebulised dornase alfa (BAC+DA) or BAC alone. On Day 1 a baseline sample was collected. From Day 1-7 participants randomised to active arm received 2.5mg BD nebulised dornase alfa in addition to BAC. In all cases BAC included dexamethasone (6 mg/day) for 10 days or until discharge, whichever was shorter as per RECOVERY (Group et al., 2021). Participants received additional treatments at the discretion of their physicians. The primary analysis was performed on samples up to Day 7. The final trial visit occurred at day 35.

Added in the revised version: Trial design and Consort diagram

Removed in the revised version: Prespecified Primary and Secondary endpoints.

A. COVASE Trial Design.

Earlier version: B. Consort diagram. Numbers not in parentheses indicate those in the intention-to-treat population, numbers in parentheses indicate the numbers in the per-protocol population.

Revised version: B. Consort diagram summary. Numbers not in parentheses indicate the participants in the intention-to-treat population and the numbers in parentheses indicate the number of participants in the per-protocol population. A complete consort diagram is shown in Supplementary figure 1.

Added in the revised version: In addition to the recruited randomised group, we included contemporary controls (CC) to increase the sample size due to practical difficulties in recruiting patients in the UK. For every COVASE participant randomised to active treatment, 2 matched CC were included. CC were admitted to UCLH over the same time as randomised patients and treated with the same BAC, including dexamethasone. CCs fulfilled the inclusion and exclusion criteria of the COVASE study and were matched for age, gender, BMI, comorbidities and CRP. CRP matching was based either on admission CRP (CC1), or matched based on first CRP after starting dexamethasone (CC2).

Added in the revised version: Eligibility

Added in the revised version: Trial participant inclusion criteria

Added in the revised version: Male and female participants, aged ≥ 18 years Dornase alfa in COVID-19: Sponsor code: 132333 Protocol Version 2.0, 25.4.2020 Page 26 of 51

Added in the revised version: Participants who are hospitalised for suspected Coronavirus (SARS-CoV)-2 infection confirmed by polymerase chain reaction (PCR) test or radiological confirmation with chest CT

Added in the revised version: Participants with stable oxygen saturation (>=94%) on supplementary oxygen

Added in the revised version: CRP >= 30 mg/L

Added in the revised version: Participants will have given their written informed consent to participate in the study and are able to comply with instructions and nebuliser

Added in the revised version: Exclusion criteria

Added in the revised version: Females who are pregnant, planning pregnancy or breastfeeding

Added in the revised version: Concurrent and/or recent involvement in other research or use of another experimental investigational medicinal product that is likely to interfere with the study medication within (specify time period e.g. last 3 months) of study enrolment

Added in the revised version: Serious condition meeting one of the following:

Added in the revised version: Respiratory distress with respiratory rate >=40 breaths/min

Added in the revised version: oxygen saturation<=93% on high-flow oxygen

Added in the revised version: Require mechanical invasive or non-invasive ventilation at screening

Added in the revised version: Concurrent severe respiratory disease such as asthma, COPD and/or ILD

Added in the revised version: Any major disorder that in the opinion of the Investigator would interfere with the evaluation of the results or constitute a health risk for the trial participant

Added in the revised version: Terminal disease and life expectancy <12 months without COVID-19

Added in the revised version: Known allergies to dornase alfa and excipients

Added in the revised version: Participants who are unable to inhale or exhale orally throughout the entire nebulisation period

Added in the revised version: So briefly

Added in the revised version: Patients were excluded if they were:

Added in the revised version: 1. pregnant, planning pregnancy or breastfeeding

Added in the revised version: 2. Serious condition meeting one of the following:

Added in the revised version: Respiratory distress with respiratory rate >=40 breaths/min

Added in the revised version: oxygen saturation<=93% on high-flow oxygen

Added in the revised version: 4. Require ventilation at screening

Added in the revised version: 5. Concurrent severe respiratory disease such as asthma, COPD and/or ILD

Added in the revised version: 6. Terminal disease and life expectancy <12 months without COVID-19

Added in the revised version: 8. Known allergies to dornase alfa and excipients

Added in the revised version: 9. Participants who are unable to inhale or exhale orally throughout the entire nebulisation period”

Added in the revised version: Participants

Added in the revised version: Adults (≥18 years of age) admitted to UCLH with confirmed SARS-Cov2 infection by RT-PCR and radiologically confirmed COVID-19 pneumonia on chest radiograph or CT-scan; oxygen saturation < 94% requiring supplemental oxygen; and evidence of hyperinflammation (CRP ≥ 30 mg/L, after administration of dexamethasone) were eligible. While all participants had CRP > 30 mg/L at screening, on two occasions the baseline CRP concentrations fell below

Added in the revised version: 30 mg/L after the patients had already agreed to participate. For full inclusion and exclusion criteria see Protocol: Supplementary Appendix 1.

Removed in the revised version: C. (left panel) Natural log CRP in BAC (HC and randomised participants; blue). (Right panel) Natural log CRP in participants randomised to BAC+DA (pink).

Removed in the revised version: D. Fitted mean (95% confidence interval) from mixed model, with natural log (CRP) over 7 days follow-up as the outcome, adjusting for natural log baseline CRP, age, sex, BMI, serious comorbidity (Diabetes, Cardiovascular disease or hypertension), time and a treatment × time interaction. P-value generated by comparing least-square means between arms. Intention to treat (ITT) population (Blue: HC and participants randomised to BAC, N=69; Pink: participants randomised to BAC+DA, N=30).

Removed in the revised version: E. Same output as in (D) but examining randomised participants only: (Blue: participants randomised to BAC, N=9; Pink: participants randomised to BAC+DA, N=30).

Removed in the revised version: Evaluations

Removed in the revised version: The baseline was defined as the last observation before the administration of dornase alfa on Day 1 for participants randomized to BAC+ dornase alfa, or the first observation after the first dose of dexamethasone for participants randomized to BAC and for HC. The participants’ clinical status was assessed each day for requirement for supplementary oxygen, admission to the intensive care unit (ICU), ventilation or high flow oxygen, and standard clinical measurements: respiration rate, oxygen saturation, systolic blood pressure, pulse rate, level of consciousness, and temperature. On day 1, and alternate days thereafter until day 7 or discharge, whichever was sooner, blood was collected for routine analysis of CRP, clinical laboratory parameters and exploratory endpoints. Participants were followed until discharge or death and/or at a follow up of at least 28 days after the last treatment day (Day 35).

Outcomes

Added in the revised version: The primary outcome was the least square (LS) mean CRP up to 7 days or at hospital discharge whichever was sooner. Data beyond 7 days was included in the model to estimate the slope. Subsequently, the estimated slope in each group was used in order to calculate the LS mean at 7 days. Pre-specified secondary outcomes included days on oxygen; time to hospital discharge; mortality by day 35; and changes in clinically relevant biomarkers including lymphocyte count and D-dimer levels. Day 35 was chosen as being likely to include most early mortality due to COVID-19 being 4 weeks after completion of a week of treatment. (i.e. d7 of treatment +28 (4 x 7 days)).

Added in the revised version: Efficacy assessments of primary and secondary outcomes in the intention-to-treat (ITT) population were performed on all randomised participants who had received at least one dose of dornase alfa if randomized to treatment. For full details see Statistical Analysis Plan. The ITT was adjusted to mitigate the following protocol violations where one participant in the BAC arm and one in the DA arm withdrew before they received treatment and provided only a baseline CRP measurement available. The participant in the DA arm was replaced with an additional recruited patient. Exploratory endpoints were only available in randomised participants and not in the CC. In this case, a post hoc within group analysis was conducted to compare baseline and post-baseline measurements.

Added in the revised version: Sample size calculation

Added in the revised version: Size calculations were produced using the proc power function in SAS Version 9.4. These were conducted to achieve 80% power to detect difference in the active arm versus the control group at 5% level of significance. Based on a mean of 99mg/L in the control group and a common standard deviation of 62mg/L derived from the literature (Han et al., 2020; Zhou, 2020), a total sample size of 90 participants would provide sufficient power to detect a greater than a 40% relative difference for CRP in the dornase alfa group compared to the control group.

Added in the revised version: This study used existing data collected at UCLH from contemporary controls (CC) admitted with COVID-19. This gave a ratio of active versus comparator of 1:2. The required power, would result in 30 participants in the active treatment group and at least 60 in the control. An additional 10 participants were recruited as a control for exploratory objectives and to compare the characteristics of enrolled participants with CC. This gave a total of 40 participants enrolled in the study and 60 CC.

Added in the revised version: Randomisation

Added in the revised version: A closed envelope method was implemented to randomize the participants into the control and active treatment groups.

Removed in the revised version: The primary efficacy outcome was the least square mean CRP up to 7 days or at hospital discharge whichever was sooner. Pre-specified secondary efficacy outcomes included days on oxygen; time to hospital discharge; clinical status at day 7 on the ordinal scale; mortality by day 35; changes in clinically relevant biomarkers including blood lymphocyte count and D-dimer levels. Other secondary outcomes were the time until discharge alive, initiation of mechanical ventilation, or ICU transfer; and the duration of ICU stay. Adverse events were recorded according to the system organ class and preferred terms in the Medical Dictionary for Regulatory Activities, version 23.0.

Removed in the revised version: Statistical analysis

Removed in the revised version: As pre-specified in the statistical analysis plan, efficacy assessments of the primary and secondary outcomes in the modified intention-to-treat population were performed on all randomised participants and received at least one dose of dornase alfa if randomized to treatment. A sample size of 90 participants was calculated to provide a power of 80% to determine a between-group difference of >40% in the primary outcome (CRP at day 7) which we considered both achievable and clinically relevant, yielding 30 participants in the active treatment group and 60 in the control group. An additional 10 participants were randomised to BAC alone as a controls for exploratory endpoints and to compare the characteristics of enrolled participants with historical controls. In summary, 40 participants enrolled in the study (30 on BAC plus dornase alfa and 10 on BAC alone, of which 30 and 9 were evaluable) and 60 historical controls which were selected using propensity score matching, with age, sex, BMI, baseline CRP, and key comorbidities defined as one or more of hypertension, diabetes or cardiovascular disease included as covariates to ensure these characteristics were as balanced as possible between those randomised to dornase alfa and the historical controls.

Removed in the revised version: All baseline data, demographics, endpoints, safety, and tolerability were summarised overall and by treatment group and by day. In general, continuous data were summarised using the mean (standard deviation), median (1st and 3rd quartiles), minimum and maximum, and categorical data were represented as frequency counts (percentages).

Removed in the revised version: For analyses relating to the primary objective, group comparisons were performed using a repeated measures mixed model, adjusted for baseline factors and with treatment as the main effect. Prior to analysis, primary and secondary endpoints were assessed for conformance to normality assumptions and the appropriate transformation was conducted if necessary.

Removed in the revised version: A re-estimation of the sample size was carried out following an interim analysis when 12 participants had been randomised.

Removed in the revised version: Safety was assessed in the randomised population. Full details of the planned statistical analysis are presented in the Statistical Analysis Plan (SAP; Supplementary Appendix 2).

Removed in the revised version: Exploratory endpoints were only available in the randomised participants and not in the historical controls. In this case, a post hoc within group analysis was conducted to compare baseline and post-baseline measurements.

Removed in the revised version: Exploratory endpoint analysis

Removed in the revised version: Peripheral venous blood was collected into EDTA or Heparin tubes, depending on clinic availability, and layered on Histopaque 1119 (Sigma-Aldrich) and centrifuged for 20 min at 800x g. The plasma, PBMC and neutrophil layers were collected. Plasma was centrifuged and frozen in liquid nitrogen, thawed and DNA levels were measured in a fluorescence plate reader using the Quant-iT™ PicoGreen dsDNA Assay Kit (P7589,Thermofisher). Data were analysed using Microsoft Excel and Graph Pad Prism software.

Removed in the revised version: Informed consent and ethics for healthy control samples

Removed in the revised version: For the cf-DNA measurements in the blood of healthy control donors, peripheral blood was isolated from consenting healthy adult volunteers, according to approved protocols of the ethics board of the Francis Crick Institute and the Human Tissue act.

Results

Added in the revised version: Patient charateristics

Removed in the revised version: Patient characteristics

Earlier version: From May 2020-October 2021, 41 participants were randomised but 1 participant in the BAC group was discharged from hospital before a second CRP measurement (Figure 1B). This participant was excluded from all analyses, except for the safety analyses. One participant withdrew consent prior to receiving any dose of dornase alfa and was replaced and excluded from all analyses. 39 participants were included in the intention-to-treat analysis set, 30 in the BAC + dornase alfa group and 9 in the BAC group. There was one treatment discontinuation, after one dose of dornase alfa, which was the participant’s decision due to a “tingling of the mouth, cough, shortness of breath” reported after receiving dornase alfa. This participant was removed from the per protocol population. All 39 participants were followed up for 35 days or until death whichever was sooner.

Revised version: From June 2020-October 2021, 41 participants were recruited and randomised. 1 participant in BAC group was discharged before a second CRP measurement, and was excluded from all, except safety, analyses (CONSORT Diagram: Figure 1B and S1). One participant withdrew consent prior to receiving dornase alfa and was replaced and excluded from all analyses. 39 participants were included in the intention-to-treat analysis set, 30 BAC+DA and 9 BAC. One participant withdrew due to side-effects. This participant was removed from the per protocol population (PPP). All 39 participants were followed up for 35 days or death whichever was sooner. Two participants were excluded from the PPP. One from BAC as this was the only patient in whom randomisation occurred prior to dexamethasone being widely used in COVID-19 and a second participant who withdrew after one dose of dornase alfa (Figure 1B and S1). The trial ended when 40 eligible participants had been recruited, although thereafter 1 had participant in the BAC arm had to be excluded from the primary and secondary endpoint analysis.

Removed in the revised version: Two participants were excluded from the per-protocol population. One from the BAC as randomisation occurred prior to dexamethasone being widely used in the treatment of COVID-19 and they were the only participant in the analysis set not to be on dexamethasone at the start of follow-up; and a second participant who withdrew after one dose of dornase alfa due to side-effects from the medication (Figure 1B).

Removed in the revised version: Baseline characteristics were generally well balanced across groups (BAC + dornase alfa, BAC, and historical controls/BAC; Table 1). Selection of historical controls via propensity score matching was successful in ensuring the means of the characteristics included in the propensity score matching were similar to those observed in the BAC + dornase alfa arm (Table 1), as well as having similar overall distributions (Supplementary Figure 1).

Earlier version: The overall mean age was 56.8 years (mean in BAC + dornase alfa group=56.8 years, mean in BAC group=56.8 years). The percentage of males was 75.8% overall (76.7% BAC + dornase alfa group, 75.4% BAC group). The most prevalent ethnicity was “White British”, with 30.3% of participants identifying in that category overall (33.3% BAC + dornase alfa group, 29.0% BAC group). The overall mean BMI was 28.0kg/m2 (mean in BAC + dornase alfa group=27.8kg/m2, mean in BAC group=28.2kg/m2). The mean baseline CRP (post dexamethasone) as defined in the primary analysis was 100.2mg/L (mean in BAC + dornase alfa group=101.9mg/L, mean in BAC group=99.5mg/L). The overall proportion of participants with a key comorbidity, defined as one or more of hypertension, diabetes, or cardiovascular disease, was 52.5% (46.7% BAC + dornase alfa group, 55.1% BAC group).

Revised version: Baseline characteristics were well balanced across groups (Table 1). Selection of 60 matched CC participants via propensity score matching was successful in ensuring the means of the characteristics included in propensity score matching were similar (Table 1 and Figure S2). The overall mean age was 56.8 years (56.8 years BAC+DA, 56.8 years BAC). The percentage of males was 75.8% (76.7% BAC+DA, 75.4% BAC). The most prevalent ethnicity was “White British”, with 30.3% of participants identifying in that category (33.3% BAC+DA, 29.0% BAC). The mean BMI was 28.0kg/m2 (27.8kg/m2 BAC+DA, 28.2kg/m2 BAC). The mean baseline CRP (post dexamethasone) was 100.2mg/L (101.9mg/L BAC=DA, 99.5mg/L) BAC. The proportion of participants with a comorbidity, defined as one or more of hypertension, diabetes, or cardiovascular disease, was 52.5% (46.7% BAC+DA, 55.1% BAC). All randomised participants except one received dexamethasone prior to randomisation, and 48 of the total 99 participants also received remdesivir or tocilizumab in addition to dexamethasone within the first 7 days. The average duration of dexamethasone treatment prior to dornase alfa was 1.13+/-0.79 days (Table S1). The last pre-dexamethasone CRP was also similar between groups, n mean of 125.0mg/L (me128.1mg/L BAC+DA, 122.7mg/L BAC). The number of days between dexamethasone initiation and baseline was 1.2 days (0.7 days BAC+DA, 1.3 days BAC). There were imbalances noted at baseline between the groups in white blood cell count, neutrophil count, procalcitonin count and D-dimer (Table S2).

Removed in the revised version: All but one (38/39) of the randomised participants received dexamethasone prior to randomisation, and 48 of the total 99 participants also received remdesivir or tocilizumab in addition to dexamethasone within the first 7 days.

Removed in the revised version: The last pre-dexamethasone CRP was also similar between groups, with an overall mean of 125.0mg/L (mean in BAC + dornase alfa group=128.1mg/L, mean in BAC group=122.7mg/L). The number of days between dexamethasone initiation and baseline was 1.2 days overall (mean in BAC + dornase alfa group=0.7 days, mean in BAC group=1.3 days).

Removed in the revised version: There were imbalances noted at baseline between the groups in white blood cell count, neutrophil count, procalcitonin count and D-dimer (Supplementary Table S1).

Clinical outcomes

Added in the revised version: Longitudinal CRP correlates with disease outcomes

Added in the revised version: To ascertain the clinical value of CRP as the COVASE trial’s primary endpoint, we first examined the correlation of CRP with mortality in a cohort of 63 hospitalized patients with severe COVID-19 (Aramburu et al., 2022b; Kurth et al., 2020; Messner et al., 2020). All participants who reached a WHO severity ordinance scale 7 were recruited at the Charité Hospital in Berlin in spring 2020 and did not receive anti-inflammatory therapies such as dexamethasone or anti-IL-6 antibody treatments. Plasma CRP measurements in 465 samples were segregated by disease outcome. This analysis revealed that deceased participants had significantly higher CRP than participants in the survivor group (Figure 2A). Clustering the 63 participant according to their longitudinal average CRP indicated that the frequency of mortality increased as the average CRP per patient increased (Figure 2B). Consistently, Mantel Cox survival analysis of the participants segregated into three groups with average longitudinal CRP ranges of 0-100 mg/L, 100-200 mg/L and 200-450 mg/L, showed that survival decreased significantly as the average CRP increased (Figure 2C). This results highlighted the strong correlation between CRP and survival in hospitalized individuals with severe COVID-19 pneumonia, suggesting that the blood CRP concentration was a strong predictor of clinical outcomes such as the probability of survival.

Added in the revised version: Longitudinal CRP predicts survival probability in severe COVID-19 pneumonia. Data from the Berlin COVID-19 study

Added in the revised version: A. Individual CRP concentrations in 465 plasmas from 63 participants with maximum WHO severity grade 7 COVID-19 pneumonia segregated into survivor and deceased groups.

Added in the revised version: B. Participants ordered by their longitudinal average CRP concentrations shown in the left column. Mortality is depicted in yellow in the right column.

Added in the revised version: C. Kaplan Meier survival probabilities (left panel) and numbers at risk (right panel) for patients segregated into three categories of longitudinal average CRP ranges: 0-100, 100-200, and 200-450 mg / L. Statistical significance (P), Hazard ratios (HR) and 95% confidence intervals (95% CI) for group 1 against group 3 and group 2 against group 3 are shown below the survival plot.

Added in the revised version: Statistics by Mann-Whitney and Mantel-Cox log rank tests.

Primary outcome

Earlier version: Individual CRP traces over time for each patient are shown in Figure 1C. Blood collection for both BAC and BAC + DA groups occurred at similar times and frequencies over the course of treatment (Supplementary Figure 2A). For the ITT group, the LS mean log (CRP) over 7 days follow-up was 3.15 (95% confidence interval [CI] 2.87 to 3.42) in the BAC + dornase alfa group (n=30), and 3.55 (95% CI, 3.35 to 3.75) in the BAC group (n=69; Table 2; Figure 1D), p=0.01. This indicates a reduction in mean CRP of approximately 33% in the BAC + dornase alfa group (23.23 mg/mL) compared to the BAC group (34.82 mg/mL) at the mean follow-up over 7 days.

Revised version: Next, we analysed the results from our COVASE participants whose individual CRP traces over time are shown in Figure S3A. Blood collection for both BAC and BAC+DA groups occurred at similar times and comparable frequencies (Figure S3B). For the ITT group, the LS mean log (CRP) over 7 days follow-up was 3.15 (95% confidence interval [CI] 2.87 to 3.42) BAC+DA (n=30), and 3.55 (95% CI, 3.35 to 3.75) in BAC (n=69) (Table 2, 3 and Figure 3A), p=0.01. This indicates a reduction in mean CRP of approximately 33% in the BAC+DA (23.23 mg/mL) compared to BAC (34.82 mg/mL) at mean follow-up over 7 days. This effect of dornase alfa on CRP was confirmed in various other subgroup analyses (Table 2): the per-modified-protocol population only; the participants who were randomised into the COVASE trial, excluding the CC (Figure 3B); participants who were randomised to BAC+DA in the COVASE trial, and CC excluding those randomised to BAC only. In addition, to ensure that the CC did not have a significantly different CRP trajectory to those randomised to BAC, we compared participants who were randomised to BAC with CC by excluding those randomised to BAC+DA and found no significant differences (Table 2). Sensitivity analyses supported the observed effect on CRP. These included the log(CRP) as an area under the curve; the CC matched for their last pre-dexamethasone CRP measurement as opposed to their first CRP after starting dexamethasone; and the effect of remdesivir or tocilizumab (Table 2).

Added in the revised version: Analysis of primary and clinical endpoints

Added in the revised version: A and B. Fitted mean (95% confidence interval) from mixed model, with natural log (CRP) over 7 days follow-up as the outcome, adjusting for natural log baseline CRP, age, sex, BMI, serious comorbidity (Diabetes, Cardiovascular disease or hypertension), time and a treatment × time interaction. P-value generated by comparing least-square means between arms. (A) Intention to treat (ITT) population (Blue: CC and participants randomised to BAC, N=69; Pink: participants randomised to BAC+DA, N=30) and (B) randomised participants only: (Blue: participants randomised to BAC, N=9; Pink: participants randomised to BAC+DA, N=30).

Added in the revised version: C. Distribution of participants based on the change in CRP measured as a ratio of the final CRP reading within the 7 day treatment period over the baseline CRP reading per patient. Statistical analysis by Fisher’s test.

Added in the revised version: D. Kaplan-Meier plot showing time to discharge from hospital from baseline. ITT population. Hazard ratio from Cox proportional hazards model adjusted for baseline CRP, age, sex, BMI, serious comorbidity (diabetes, cardiovascular disease of hypertension). P-value from log-rank test. (Blue: CC and participants randomised to BAC, N=69. Pink: participants randomised to BAC+DA, N=30).

Added in the revised version: E. Kaplan-Meier plot showing time to death over 35 days follow up. ITT population. Hazard ratio from Cox proportional hazards model adjusted for baseline CRP, age, sex, BMI, serious comorbidity (Diabetes, Cardiovascular disease of hypertension). P-value from log-rank test. (Blue: CC and participants randomised to BAC, N=69. Pink: participants randomised to BAC+DA, N=30). Abbreviations: BAC-best available care, CRP-C-reactive protein, DA-dornase alfa, ITT-intention-to-treat.

Added in the revised version: To better understand the effects of dornase alfa on the CRP of the treated population we also performed a frequency distribution analysis of the change in CRP over time. Instead of fitting a slope on the CRP values of the individual samples of each participant, we plotted the fraction between the CRP measurement of the baseline and final sample collected for each participant during the 7 day treatment period. We then plotted the distribution of participants within the different ranges of change in CRP (Figure 3C). Compared to BAC, the BAC+DA treatment group exhibited a distribution towards larger ranges in CRP reduction with better suppression of the fraction of patients that did not respond to treatment as indicated by an increase in CRP (fraction: 1 to 6). These results indicated that the addition of dornase alfa to dexamethasone cause a significant and sustained reduction in CRP in a greater number of COVID-19 infected participants.

Removed in the revised version: This effect of dornase alfa on CRP was confirmed in various other subgroup analyses and are shown in Table 2: the per-protocol population only; participants who were randomised into the COVASE trial, excluding the historical controls (Figure 1E); participants who were randomised to BAC + dornase alfa in the COVASE trial, and HC,. excluding those randomised to BAC only (Table 2).

Removed in the revised version: In addition, to ensure that the HCs did not have a significantly different CRP trajectory to those randomised to BAC, we compared participants who were randomised to BAC with HC by excluding those randomised to BAC + dornase alfa and found no significant differences (Table 2).

Removed in the revised version: Sensitivity analyses were also conducted and continued to support the observed effect on CRP. These are also shown in Table 2 and included: log(CRP) as an area under the curve; historical controls matched for their last pre-dexamethasone CRP measurement as opposed to their first CRP after starting dexamethasone; and the effect of remdesivir or tocilizumab.

Secondary outcomes

Earlier version: Length of hospitalisation was analysed as a time-to-event outcome of alive discharge from hospital censored at 35 days. The hazard ratio observed in the Cox proportional hazards model was 1.63 (95% CI, 1.01 to 2.61), p=0.03 (Table 3 and Figure 2A). Showing that throughout 35 days follow-up, there was a 63% higher chance of discharge alive at any given time-point in the BAC + dornase alfa group compared to the BAC group. Although the rate of discharge was similar in 50% of patients, 80% discharge occurred by 8 days in the dornase alfa group whereas, whereas the same proportion was reached at 30 days in the BAC group, suggesting that dornase alfa may be beneficial to patients that do not respond efficiently to dexamethasone alone. This trend was also seen when only the COVASE participants were considered, although not powered to reach significance and with a smaller HR of 1.18 (95% CI, 0.52-2.69), p=0.62, (Supplementary Table 2 and Supplementary Figure 2B).

Revised version: Length of hospitalisation was analysed as a time-to-event outcome of alive discharge from hospital censored at 35 days. The hazard ratio observed in the Cox proportional hazards model was 1.63 (95% CI, 1.01 to 2.61), p=0.03 (Table 4, Figure 3D and S3C). Showing that throughout 35 days follow-up, there was a 63% higher chance of discharge alive at any given time-point in BAC+DA compared to BAC. Although the rate of discharge was similar in 50% of patients, 80% discharge occurred by 8 days in BAC+DA whereas, whereas the same proportion was reached at 30 days in BAC. This trend was also seen when only the randomised participants were considered, although not powered to reach significance and with a smaller HR of 1.18 (95% CI, 0.52-2.69), p=0.62, (Table S3 and Figure S3D).

Removed in the revised version: Analysis of secondary endpoints and exploratory endpoints.

Removed in the revised version: A. Kaplan-Meier plot showing time to discharge from hospital from baseline. ITT population. Hazard ratio from Cox proportional hazards model adjusted for baseline CRP, age, sex, BMI, serious comorbidity (Diabetes, Cardiovascular disease of hypertension). P-value from log-rank test. (Blue: HC and participants randomised to BAC, N=69. Pink: participants randomised to BAC+DA, N=30).

Removed in the revised version: B. Kaplan-Meier plot showing time to death over 35 days follow up. ITT population. Hazard ratio from Cox proportional hazards model adjusted for baseline CRP, age, sex, BMI, serious comorbidity (Diabetes, Cardiovascular disease of hypertension). P-value from log-rank test. (Blue: HC and participants randomised to BAC, N=69. Pink: participants randomised to BAC+DA, N=30). Abbreviations: BAC-best available care, CRP-C-reactive protein, DA-dornase alfa, ITT-intention-to-treat.

Removed in the revised version: C. Difference between the lymphocyte count for each day of the treatment period and the baseline in each patient who exhibited lymphopenia at baseline (<1×109 lymphocytes/mL). Mean and 95%CI interval is shown with statistical analysis by two-way Anova.

Removed in the revised version: D. Mean D-dimer levels per day in randomised BAC (blue) and BAC+DA (pink) patients with error bars depicting 95% CI in randomised BAC (blue) and BAC+DA (DA) patients (pink). Statistical difference by by mixed effects Anova analysis.

Removed in the revised version: E. Mean cf-DNA levels per day in randomised BAC (blue) and BAC+DA (pink) patients, with error bars depicting standard deviation. Statistical analysis by mixed effects Anova.

Removed in the revised version: F. Correlation between the final cf-DNA levels and ratio of CRP at day-7 normalized to the baseline CRP (CRPfinal/CRPbaseline) per patient. Fitting by non-linear regression.

Earlier version: Over 7 days of follow up there was no significant difference between BAC + dornase alfa versus BAC alone in either the fraction of participants admitted to ICU (23.3% versus 21.74%), p=0.866, or the length of ICU stay, least squares mean 21.25 (95% CI, 4.65 to 37.84) hours versus 19.85 (95% CI, 8.00 to 31.70) hours, p=0.883. The same was seen over 35-day follow-up, with least squares mean 55.21 95% CI, -23.59 to 134.00) hours versus 60.60 (95% CI, 4.34 to 116.86) hours, p=0.905. At any point during the 35 days follow-up, 23% of the BAC + dornase alfa group were admitted to ICU compared to 23.19% in the BAC group, p= 0.983 (Supplementary Table 3).

Revised version: Over 7 days of follow up there was no significant difference between BAC+DA versus BAC alone in either the fraction of participants admitted to ICU (23.3% versus 21.74%), p=0.866, or the length of ICU stay, LS mean 21.25 (95% CI, 4.65 to 37.84) hours versus 19.85 (95% CI, 8.00 to 31.70) hours, p=0.883. The same was seen over 35-day follow-up, LS mean 55.21 95% CI, −23.59 to 134.00) hours versus 60.60 (95% CI, 4.34 to 116.86) hours, p=0.905. At any point during the 35 days follow-up, 23% of BAC+DA were admitted to ICU compared to 23.19% BAC, p= 0.983 (Table S4).

Earlier version: There was no significant difference in time requiring oxygen between the two groups, at either 7 days, least squares mean 94.32 (95% CI, 72.8 to 115.79) hours, versus 88.96 (95% CI, 73.64 to 104.29) hours, p=0.662, or 35 days, least squares mean 133.22 (95% CI, 52.01 to 214.43) hours versus 156.35 (95% CI, 98.36, 214.33) hours, p=0.618. At 35 day follow up if we look at COVASE participants, there are only 9 participants to evaluate, but mean oxygen use tends to a reduction of 123 hours with BAC + dornase alfa, versus 241 hours for BAC, p=0.187 (Supplementary Table 3)

Revised version: There was no significant difference in time requiring oxygen between the two groups, at either 7 days, LS mean 94.32 (95% CI, 72.8 to 115.79) hours, versus 88.96 (95% CI, 73.64 to 104.29) hours, p=0.662, or 35 days, LS mean 133.22 (95% CI, 52.01 to 214.43) hours versus 156.35 (95% CI, 98.36, 214.33) hours, p=0.618. At 35 days follow up, there were only 9 randomised participants to evaluate, but mean oxygen use tends to a reduction of 123 hours with BAC+DA, versus 241 hours for BAC, p=0.187 (Table S4)

Earlier version: Over 35 days follow up, 1 person amongst the 30 patients in the BAC + dornase alfa group died, compared to 8 of the 69 participants in the BAC group. The hazard ratio observed in the Cox proportional hazards model was 0.47 (95% CI, 0.06 to 3.86), indicating a trend towards a reduced chance of death at any given time-point in the BAC + dornase alfa group compared to the BAC group, but this did not reach significance p= 0.460 (Figure 2B).

Revised version: The time to event data was censored at 28 days post last dose (up to d35) for the randomised participants and at the date of the last electronic record for the CC. Over 35 days follow up, 1 person amongst the 30 patients in BAC+DA died, compared to 8 of the 69 participants in BAC. The hazard ratio observed in the Cox proportional hazards model was 0.47 (95% CI, 0.06 to 3.86), indicating a trend towards a reduced chance of death at any given time-point in BAC+DA compared to BAC, but this did not reach significance p= 0.460 (Figure 3E and S3D). The hazard ratio observed in the Cox proportional hazards model (95% CI) was 0.47 (0.06, 3.86), which estimates that throughout 35 days follow-up, there was a 53% reduced chance of death at any given timepoint in the BAC + dornase-alfa group compared to the BAC group, though the confidence intervals are wide due to a small number of events. The p-value from a log-rank test was 0.460, which does not reach statistical significance at an alpha of 0.05.

Earlier version: There was no significant difference at either 7- or 35-days follow-up, in the number of participants that required mechanical ventilation in the BAC + dornase alfa group compared with the BAC group (16.67% vs 13.04%), p=0.628. Amongst participants that were ventilated, the mean length of mechanical ventilation at 7 days follow-up in the BAC + dornase alfa group was 76.8 hours, compared to 88.78 in the BAC group. At 35 days follow-up, the mean length of mechanical ventilation in the BAC + dornase alfa group was 76.8, compared to 411.17 in the BAC group (Supplementary Table 3).

Revised version: There was no significant difference at either 7-or 35-days follow-up, in the number of participants that required mechanical ventilation in BAC+DA compared with BAC (16.67% vs 13.04%), p=0.628. Amongst participants that were ventilated, the mean length of mechanical ventilation at 7 days follow-up in BAC+DA was 76.8 hours, compared to 88.78 in BAC. At 35 days follow-up, the mean length of mechanical ventilation in BAC+DA was 76.8 hours compared to 411.17 hours in BAC (Table S4). There was no significant difference in superadded bacterial pneumonia at either 7-or 35-days follow-up: 7 days, 1 (3.33%) participant in BAC+DA compared to 3 (4.35%) participants in BAC, p= 0.934; 35 days, 2 (6.67%) participants in BAC+DA had bacterial pneumonia, compared to 3 (4.35%) participants in BAC, p=0.548 (Table S4).

Removed in the revised version: There was no significant difference in superadded bacterial pneumonia at either 7- or 35-days follow-up: 7 days, 1 (3.33%) participant in the BAC + dornase alfa group compared to 3 (4.35%) participants in the BAC group, p= 0.934; 35 days, 2 (6.67%) participants in the BAC + dornase alfa group had bacterial pneumonia, compared to 3 (4.35%) participants in the BAC group, p=0.548 (Supplementary Table 3).

Removed in the revised version: Blood analysis with no adjustment for multiple testing showed a significant treatment effect in BAC + dornase alfa group vs. BAC group for three parameters: lymphocyte counts, D-dimer, a marker of coagulation and procalcitonin (PCT).

Earlier version: First, the dornase alfa treated group exhibited higher lymphocyte counts with a least-squares mean of 0.87 (95% CI, 0.76-0.98) in the BAC group vs. 1.08 (95% CI, 0.92-1.27) in the BAC + dornase alfa group, p=0.02 (Table 3 and Supplementary Table 2). In particular, patients with lymphopenia at baseline (<1×109 lymphocytes/L) exhibited a greater increase in blood lymphocyte numbers in the BAC + dornase alfa group than in the BAC group during the entire length of treatment (Figure 2C).

Revised version: Blood analysis with no adjustment for multiple testing showed a significant treatment effect in BAC+DA vs. BAC group for three parameters: lymphocyte counts, D-dimer, and procalcitonin (PCT). First, BAC+DA exhibited higher lymphocyte counts with LS mean of 0.87 (95% CI, 0.76-0.98) in the BAC group vs. 1.08 (95% CI, 0.92-1.27) in BAC+DA, p=0.02 (Table 4 and Table S3). Patients with lymphopenia at baseline (<1×109 lymphocytes/L) exhibited a greater increase in blood lymphocyte numbers in BAC+DA than in BAC during the entire length of treatment (Figure 4A). Furthermore, D-dimer levels were lower in BAC+DA compared to BAC, with LS mean D-dimer difference of 1657 (95% CI, 3131-877) (Table 4, Table S3, Figure 4B and Figure S4A). Our analysis indicated lower levels of PCT in participants that received BAC+DA, mean 0.18 ng/mL (95% CI, −0.2-0.56) compared to BAC alone, mean 1.31 ng/mL (95% CI, 0.56-2.05), p=0.005 (Table 4). Repeat analysis excluding the CC population replicated these results, and changes in all 3 parameters were significant (Table S3).

Added in the revised version: Analysis of secondary and exploratory endpoints in blood

Added in the revised version: A. Difference between the lymphocyte count for each day of the treatment period and the baseline in each patient who exhibited lymphopenia at baseline (<1×109 lymphocytes/mL). Mean and 95%CI interval is shown with statistical analysis by two-way Anova.

Added in the revised version: B. Mean D-dimer levels per day in randomised BAC (blue) and BAC+DA (pink) participants with error bars depicting 95% CI in randomised BAC (blue) and BAC+DA (DA) participants (pink). Statistical difference by by mixed effects Anova analysis.

Added in the revised version: C. Mean cf-DNA levels per day in randomised BAC (blue) and BAC+DA (pink) participants, with error bars depicting standard deviation. Statistical analysis by mixed effects Anova.

Added in the revised version: D. Correlation between the final cf-DNA levels and ratio of CRP at day-7 normalized to the baseline CRP (CRPfinal/CRPbaseline) per participant. Fitting by non-linear regression.

Added in the revised version: E. Correlation between D-dimer and cell-free (cf) DNA levels in the blood of participants randomised to BAC (blue) or to BAC+DA (DA) (pink), where samples have been segregated depending on whether the corresponding levels of cf-DNA were below or above 100 μg/mL. Statistical analysis by unpaired parametric t-test.

Removed in the revised version: Furthermore, D-dimer levels were lower in the BAC + dornase alfa group compared to the BAC group, with a least-squares mean D-dimer difference of 1657 (95% CI, 3131-877) (Table 3, Supplementary Table 2, Figure 2D and Supplementary Figure 2C).

Removed in the revised version: Procalcitonin (PCT) is marker of bacterial infection that is also elevated in many patients with severe COVID-19 infection. Our analysis indicated lower levels of PCT in patients that received dornase alfa, with a mean value of 0.18 ng/mL (95% CI, -0.2-0.56) compared to those treated with BAC alone with a mean of 1.31 ng/mL (95% CI, 0.56-2.05), p=0.005 (Table 3). On repeat analysis excluding the historical control population the results were replicated and changes in these 3 parameters reached statistical significance (Supplementary table 2).

Exploratory outcomes

Earlier version: Given the role of circulating cf-DNA in pathology, we examined whether the pulmonary administration of dornase alfa influenced systemic cf-DNA levels in plasma. There was no difference in baseline plasma cf-DNA levels on Day 1 between the two groups. However, during the treatment period cf-DNA was reduced in participants randomised to BAC+ dornase alfa compared to BAC alone (Figure 2E and Supplementary Figure 2D). We also examined whether cf-DNA levels correlated with D-dimer and CRP levels. Samples that contained cf-DNA above 100μg/mL exhibited significant higher D-dimer levels compared to samples containing cf-DNA levels below 100μg/mL (Supplementary Figure 2E). Moreover, there was a positive correlation between the levels of cf-DNA in the final sample collected during the treatment period and the ratio of final to baseline CRP (CRPfinal/CRPbaseline) indicating that changes in CRP were inversely proportional to the final cf-DNA levels in all patients independently of treatment (Figure 2F).

Revised version: Given the role of circulating cf-DNA in pathology, we examined whether the pulmonary administration of dornase alfa influenced systemic cf-DNA levels. Compared to the plasma of anonymized healthy donor volunteers at the Francis Crick institute (HD), plasma cf-DNA levels were elevated in both BAC and DA-treated COVASE participants. There was no difference in baseline plasma cf-DNA levels between the two groups. However, during the treatment period cf-DNA was reduced in participants randomised to BAC+DA compared to BAC alone (Figure 4C and Figure S4B). There was a positive correlation between the levels of cf-DNA in the final sample collected during the treatment period and the ratio of final/baseline CRP (CRPfinal/CRPbaseline) (Figure 4D). Moreover, samples with cf-DNA >100μg/mL had significantly higher D-dimer levels compared to samples with cf-DNA levels <100μg/mL (Figure 4E). Hence, dornase alfa treatment was associated with a suppression of circulating cf-DNA that correlated with a more sustained reduction in CRP.

Safety

Earlier version: Dornase alfa was very well tolerated, with no systemic effects and this was consistent with its short half-life and lack of systemic exposure. There were 10 reported AEs by 9 participants in the randomised BAC arm versus 30 AEs reported by 30 participants in the BAC + dornase alfa arm. Of these, one was reported by the clinical team as definitely related to the study drug and one as unlikely to be related to the study drug (Supplementary Table S4). ‘Tingling of the mouth’ after using the nebuliser was attributable to the drug whilst ‘headache’ was unlikely to be related to the study drug. The AE data reflect the clinical trial and post-marketing experience of using Pulmozyme at the recommended dose regimen. Adverse reactions attributed to Pulmozyme are reported as rare (< 1/1000). No treatment-related serious adverse events (SAEs) occurred in any participants.

Revised version: Dornase alfa was well tolerated with no systemic effects, consistent with its short half-life (Table S5). There were 10 reported adverse events (AEs) in 9 participants with BAC versus 30 AEs in 30 participants with BAC+DA (Table S6). Of these, one was reported by the clinical team as definitely related and one as unlikely to be related to study drug (Table S5). No treatment-related serious AEs were reported.

Discussion

Added in the revised version: Our findings show that nebulised dornase alfa significantly reduced CRP in patients with severe COVID-19 pneumonia and systemic inflammation even when already receiving dexamethasone. This was robust with several sensitivity analyses including analysis of only randomised participants (i.e excluding CC). The secondary outcomes were consistent with the same directionality. Despite not being formally powered for length of stay, dornase alfa reduced time to discharge over 35 days.

Added in the revised version: Identification of novel/repurposed treatments effective for COVID-19 were hampered by patient recruitment to competing studies during a pandemic. This resulted in small studies with inconclusive or contrary findings. Our study design took advantage of frequent repeated measures of CRP in each patient, to allow a smaller sample size to determine efficacy/futility. A similar CRP-based approach was applied in CATALYST (Fisher et al., 2022) and ATTRACT studies (Tornling et al., 2021). In addition, we used contemporary controls as additional comparators to efficiently use limited resources.

Removed in the revised version: In this Phase IIa trial involving hospitalised participants with severe COVID-19 pneumonia and systemic inflammation, we found that nebulised human recombinant DNA-ase (dornase alfa) significantly reduced CRP over 7 days, versus BAC, which included dexamethasone, alone. The finding was robust after several sensitivity analyses and no safety concerns resulted from the use of dornase alfa in this patient group. No significant mortality benefit was associated with this reduction, although this phase IIa trial was not powered for this outcome. Moreover dornase alfa reduced time to discharge over 35 days follow up. Adverse events were balanced between the two groups with no treatment related SAEs.

Earlier version: The RECOVERY trial resulted in dexamethasone becoming standard care in patients with COVID-19 pneumonitis requiring supplemental oxygen. We recruited participants dependent on a CRP ≥ 30 mg/L, on the day after receiving their first dexamethasone dose to minimise steroid-dependent effects on CRP. The finding that dornase alfa can significantly reduce CRP in participants receiving dexamethasone suggests a complementary mode of action to deliver sustained reduction in inflammation. Dornase alfa may also prove suitable for the treatment of patients with mild viral pneumonia, in contrast to dexamethasone which is only effective in patients with an oxygen requirement, and potentially harmful in milder cases17.

Revised version: Following the RECOVERY trial, dexamethasone became standard care in patients with COVID-19 pneumonitis requiring oxygen. We recruited participants with CRP ≥ 30 mg/L, on the day after receiving dexamethasone to minimise steroid-dependent effects on CRP. The finding that dornase alfa can significantly reduce CRP (by average 33%) in participants receiving dexamethasone suggests a complementary mechanism of action. Dornase alfa may provide a treatment choice for patients with mild COVID-19 pneumonia not requiring oxygen, in whom dexamethasone may be harmful(Group et al., 2021).

Added in the revised version: Our primary effect was consistently reflected in secondary outcomes. Dornase alfa increased the chance of live discharge by 63% at any time up to 35 days, with a trend to reduce hospital stay. This impact on hospital occupancy during a pandemic cannot be overestimated. In addition, dornase alfa significantly increased lymphocyte counts and reduced D-dimer. CRP levels are associated with venous thromboembolic disease in COVID-19 with the worst outcomes seen in patients with high CRP and D-dimers (Smilowitz et al., 2021). In a metanalysis of 32 studies involving 10 491 COVID-19 patients, elevated CRP (OR 4.37, p<0.00001), lymphopenia (3.33, p<0.00001), elevated D-dimer (3.39 (p<0.00001) and elevated PCT (6.33, p<0.00001) were independent markers of poor outcomes(Malik et al., 2021).

Added in the revised version: Our study design offered a solution to the early screening of compounds for inclusion in larger platform trials. The study took advantage of frequent repeated measures of quantifiable CRP in each patient, to allow a smaller sample size to determine efficacy/futility than if powered on clinical outcomes. We applied a CRP-based approach that was similar to the CATALYST and ATTRACT studies (Fisher et al., 2022; Tornling et al., 2021). CATALYST showed in much smaller groups (usual care, 54, namilumab, 57 and infliximab, 35) that namilumab that is an antibody that blocks the cytokine GM-CSF reduced CRP even in participants treated with dexamethasone whereas infliximab that targets TNF-α had no significant effect on CRP. This led to a suggestion that namilumab should be considered as an agent to be prioritised for further investigation in the RECOVERY trial. A direct comparison of our results with CATALYST is difficult due to the different nature of the modelling employed in the two studies. In general, Dornase alfa exhibited comparable significance in the reduction in CRP compared to standard of care as described for namilumab at a fraction of the cost. Furthermore, endonuclease therapies may prove superior to cytokine blocking monotherapies, as they are unlikely to increase the risk for microbial co-infections that have been reported for antibody therapies that neutralize cytokines that are critical for immune defence such as IL-1β, IL-6 or GM-CSF.

Added in the revised version: In addition, Dornase alfa reduced circulating cf-DNA levels, suggesting that this nebulized enzyme exerts systemic effects on circulating chromatin and inflammatory mediators. The reduction in cf-DNA suggests that by stripping the DNA from chromatin, dornase alfa suppresses the proinflammatory properties of histones and potentiates their degradation(Aramburu et al., 2022a; Papayannopoulos et al., 2011). The inverse correlation between final circulating cf-DNA and change in CRP during treatment highlights the link between circulating chromatin and systemic inflammation. The correlation of d-dimer with cf-DNA and the reduction in d-dimer with dornase alfa is consistent with the reported pro-thrombotic role of NETs in the alveoli (Radermecker et al., 2020). Moreover, the difference in the recovery rates for lymphopenia is in line with a link between extracellular chromatin and lymphocyte death in sepsis (Ioannou et al., 2021). Another reason why recombinant endonuclease treatment may be warranted in patients with acute infections is the reduction in endogenous chromatin degradation capacity that correlates with a high risk for mortality (Aramburu et al., 2022a).

Removed in the revised version: Given the lack of support for COVID-19 studies outside approved UK platform trials, we employed historical controls to allow us to complete the study more quickly and match patients more closely. Still these historical controls include patients that had been admitted on days when competing studies were unable to recruit because of clinical trials pharmacy closures, or patients that had been recruited to the BAC arm of other clinical trials. The historical controls all met the inclusion criteria for COVASE and did not meet any exclusion criteria. They were matched as closely as possible for key parameters and comorbidities. Despite this potential heterogeneity we were still able to show a reduction in CRP and a tendency to reduce hospital stay in dornase alfa treated participants. Our trial was not powered to overcome confounders, such as the use of antivirals and tocilizumab, an IL-6 inhibitor, recognised to reduce CRP18. Although underpowered, we demonstrate a trend to a reduction in CRP with dornase alfa in participants that had received tocilizumab and/or remdesivir, and those that had not.

Removed in the revised version: By stripping the DNA from chromatin, dornase alfa suppresses the proinflammatory properties of histones and potentiating their degradation by serum and NET proteases14,19. Suprisingly, dornase alfa reduced cf-DNA levels in the circulation athough the treatment is only localized to the lung due to the short half-life of the enzyme in the bloodstream. Hence, decreases in blood cf-DNA levels likely reflect the enhanced local clearance of pulmonary NETs and DNA released by other cell types. Consistently, we noted an inverse correlation between the final readings of circulating cf-DNA and the magnitude of decrease in CRP within the 7-day treatment period, further supporting a functional link between chromatin levels and systemic inflammation in COVID-19 pneumonitis. The reduction in D-dimer is also consistent with the pro-thrombotic role of NETs in the alveoli of these patients3. Consistently, CRP and D-dimer correlated with cf-DNA in all patients. Moreover, the increased recovery of lymphopenia is consistent with our recent finding that cf-chromatin promotes lymphocyte death in sepsis8. By countering these pathogenic properties of cf-chromatin, dornase alfa counters coagulopathy and lymphopenia in these patients. Recombinant DNAse treatment may also be beneficial due to a significant number of SARS-CoV-2 and microbial sepsis patients exhibiting defects in plasma cf-DNA degradation19 which is a critical factor in patient survival.

Earlier version: Whilst immunisation has greatly reduced the numbers of patients admitted to hospitals with COVID-19 pneumonia there is still a need for virally agnostic therapies that retain efficacy even as viruses mutate. Moreover, nebulised dornase alfa can be safely administered outside the health-care setting. Three other trials of dornase alfa in patients with COVID-19 have data published in peer reviewed manuscripts reporting improvements in oxygen requirements20-22 (Holliday et al., 2021; Okur et al., 2020; Weber et al., 2020). One small study indicated improvements in the plasma and sputum proteomic profile (Fisher et al., 2021). However, these studies examined a small number of patients and the study designs, patient populations and endpoints differ from those used here.

Revised version: Nebulized dornase alfa treatment has several advantages. Whilst immunisation has reduced COVID-19 hospital admissions, a substantial population remain unvaccinated or immunocompromised. Hence, there will always be a need for virally agnostic therapies that retain efficacy as viruses mutate. Nebulised dornase alfa controls immune pathology in a range of infections. Moreover, it can be administered safely and effectively outside the health-care setting. Three other small trials of dornase alfa (totalling 18 patients) in COVID-19 have reported improved oxygenation (Holliday et al., 2021; Okur et al., 2020; Weber et al., 2020). One small study indicated improvements in plasma and sputum proteomic profiles (Fisher et al., 2021). Despite differing study designs, patient populations and endpoints, the overall trend suggests an improvement in clinical outcomes.

Added in the revised version: There are limitations in this study. This single centre open-label study was designed to powerfully report futility or efficacy despite including just under 40 patients and 60 CC. However, the trial was not powered to report mortality nor to overcome confounders, such as the use of antivirals and tocilizumab, an IL-6 inhibitor recognised to reduce CRP (Galvan-Roman et al., 2021), or the impact of an open-label study on influencing other therapeutic/ discharge decisions. Although underpowered, we demonstrate a trend to a reduction in CRP with dornase alfa in participants that had received tocilizumab and/or remdesivir, and those that had not. Moreover, the open-label nature of the study could potentially introduce bias. Nevertheless, we tried to minimize this by applying a standard testing schedule and discharge criteria.

Added in the revised version: In conclusion, we demonstrate that nebulised dornase alfa significantly reduces inflammation in hospitalised patients with severe COVID-19 pneumonia leading to improved clinical profiles and earlier discharge from hospital. These results confirm the role of cf-DNA in promoting systemic inflammation, coagulopathy and immune dysfunction in acute respiratory infections and suggest that recombinant DNases provide an additional mode of anti-inflammatory action that can broaden the effectiveness and sustainability of anti-inflammatory regimens. These encouraging data warrant further investigation in other respiratory infections characterized by hyperinflammation in the setting of cell lysis and/ or NET formation.

Removed in the revised version: In conclusion, we have demonstrated that nebulised dornase alfa significantly reduces inflammation in hospitalised patients with severe COVID-19 pneumonia with a trend towards clinical benefit with reduced oxygen requirements and earlier discharge from hospital in patients that received this treatment. These very encouraging preliminary findings warrant further investigation in larger studies.

Supporting information